# SECTION I - CODING TEST


## Part 1 - Numpy and ML Basics

```python
import numpy as np
```

### Q Samples and Features
Given a dataset in the form of a numpy array where axis = 0 indexes samples and axis = 1 indexes features. Write a function that returns number of samples in dataset and number of samples in each feature.

<!-- grade_id: cell-5e96ed30afc15a3e, points: 0, grade: false, solution: true -->
```python
def samps_feats(x):
    """
    Inputs:
        x: numpy array of shape (samples, features)
    Outputs:
        s: int, number of samples in x
        f: int, number of features in x
    """
    # YOUR CODE HERE
    return s, f
```

<!-- grade_id: cell-bbe2be4b7abbd82a, points: 1, grade: true, solution: false -->
```python
"""Testing"""
assert samps_feats(np.arange(12).reshape(4, 3)) == (4, 3)
print('Sample Test passed', '\U0001F44D')
```

### Q Third Column

<!-- grade_id: cell-592b835e0cfd264a, points: 0, grade: false, solution: true -->
```python
def thirdColumn(a):
  '''
  Given a two dimensional array a of shape (m, n) m>=3, n>=3. Return the third column of a.
  
  Input: 
      a: Numpy array of shape (m, n)
  Output:
      c:  Numpy array of shape (m,)
     
  '''
  # YOUR CODE HERE
```

<!-- grade_id: cell-42d7a633a168e12f, points: 1, grade: true, solution: false -->
```python
"""Tests"""
a = np.arange(9).reshape(3, 3)
assert np.all(thirdColumn(a) == np.array([2, 5, 8]))
print('Test passed', '\U0001F44D')


```

### Q Precision
Given a confusion matrix $cm$ of shape ($N_y$, $N_y$) and a class $cls$ where $0<=cls<=N_y-1$. 
Compute and return the precision for that class from the confusion matrix 



<!-- grade_id: cell-cabf3cdd9585e5bd, points: 0, grade: false, solution: true -->
```python
def precision(cm, cls):
    """
    Inputs:
        cm : np array of shape (Ny, Ny)
        cls : int, 0<=cls<=Ny-1
    Outputs:
        prec : float, precision of class cls in cm
    """
    # YOUR CODE HERE
    return prec
```

<!-- grade_id: cell-760a0d18c030c768, points: 2, grade: true, solution: false -->
```python
"""Test for precision"""
cm = np.arange(9).reshape(3, 3)
assert np.isclose(precision(cm, 2), 0.5333333333, atol = 0.001) 
print('Sample Test passed', '\U0001F44D')
```

### Q Sigmoid derivative
Given an input $z$. Write a function that returns the derivative of the sigmoid($\sigma$) function.<br>
$\sigma(z) = \frac{1}{1+e^z}$

<!-- grade_id: cell-e0cab6ca0f4ccc5c, points: 0, grade: false, solution: true -->
```python
def sigmoid_derivative(z):
    """
    Inputs:
        z: float, input 
    Outputs:
        sig_der: float, derivative of sigmoid function at z
    """
    # YOUR CODE HERE
    return sig_der
```

<!-- grade_id: cell-48064ee9114778b4, points: 2, grade: true, solution: false -->
```python
assert np.isclose(sigmoid_derivative(99999), 0, atol = 0.0001)
print("Sample Tests passed", '\U0001F44D')
```

### Q   Find mean value in each column
Find mean value in each column in a numpy 2-d array and return the 1-d array


<!-- grade_id: cell-c70acfa1ba68a9e8, points: 0, grade: false, solution: true -->
```python
def find_mean_column_wise(a):
  '''  
  Input:
      a: numpy array of shape (m, n)
  Output:
      b: numpy array of shape (n, )
  
  '''
  # YOUR CODE HERE
  return b
```

<!-- grade_id: cell-f77fbcfbf0e4d1e3, points: 3, grade: true, solution: false -->
```python
"""Testing"""
a = find_mean_column_wise(np.array([[9, 6, 4],[8, 6, 1],[5, 3, 6],[3, 3, 3],[2, 1, 7]]))
assert np.all(np.isclose(a, np.array([5.4, 3.8, 4.2])))
print('Test passed', '\U0001F44D')
```

## Part 2 - Keras with MLN (Heart Disease Dataset)

#### Import libraries and Data

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import requests
```

```python
### Downloading data from google drive. You need not mess with this code.

import requests

def download_file_from_google_drive(id, destination):
    URL = "https://docs.google.com/uc?export=download"

    session = requests.Session()

    response = session.get(URL, params = { 'id' : id }, stream = True)
    token = get_confirm_token(response)

    if token:
        params = { 'id' : id, 'confirm' : token }
        response = session.get(URL, params = params, stream = True)

    save_response_content(response, destination)    

def get_confirm_token(response):
    for key, value in response.cookies.items():
        if key.startswith('download_warning'):
            return value
    return None

def save_response_content(response, destination):
    CHUNK_SIZE = 32768

    with open(destination, "wb") as f:
        for chunk in response.iter_content(CHUNK_SIZE):
            if chunk: # filter out keep-alive new chunks
                f.write(chunk)
                
if __name__ == "__main__":
    file_id = '1vuqmYrgmXD5NE2-mN3qNrsg8l1wCaiif'
    destination = 'heart.csv'
    download_file_from_google_drive(file_id, destination)
```

#### Segregating input and output data


```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import requests


# Importing and cleaning data using pandas library
data = pd.read_csv('heart.csv')

## Last column is output features. Rest are inputs
X = data.iloc[:, 0:-1].values

# Last column is output labels
y = data.iloc[:, -1].values
```

### Question 1
#### Visualizing the data output

<!-- grade_id: cell-325781143dca3b66, points: 0, grade: false, solution: true -->
```python
# find the number of classes in the output data in this classification problem
def Num_of_classes(y_raw):
    """
    Inputs:
        y_raw: numpy array of labels
    Outputs:
        Ny: number of classes
    """
    # YOUR CODE HERE
    
    return Ny
```

<!-- grade_id: cell-b3920480ab5580aa, points: 1, grade: true, solution: false -->
```python
"""Test for Num_of_classes"""
```

### Question 2
#### one-hot encoding

<!-- grade_id: cell-fdb5d153c6e6e7e8, points: 0, grade: false, solution: true -->
```python
#Use Keras to do One-hot encoding
def keras_one_hot(y_int, ny):
  '''
  Input:
      y_int: an int in {0, Ny -1 }
      ny: Number of classes, e.g., 2 here.
  Output:
      Y: a vector of shape (Ny,)
  '''
  # YOUR CODE HERE
  return Y
```

<!-- grade_id: cell-fd6e8b2a42b86195, points: 2, grade: true, solution: false -->
```python
"""Test for keras_one_hot"""
```

### Question 3
#### Create Y_o which is one-hot encoding of Y using above functions

<!-- grade_id: cell-a448645c12ef97f2, points: 0, grade: false, solution: true -->
```python
def create_Y_o(y):
    """
    Inputs:
        y: numpy array of class labels
    Outputs:
        Y_o: numpy array of shape(samples, Ny) with one-hot encodings of y
        ny: number of unique classes
    """
    
    # YOUR CODE HERE
    
    return Y_o, ny

Y_o, ny = create_Y_o(y)
```

<!-- grade_id: cell-0a7501407e4aea79, points: 2, grade: true, solution: false -->
```python
"""Test for create_Y_o"""
```

### Question 4
#### Split the data into train data and test data

<!-- grade_id: cell-90f585b9d28d09a3, points: 0, grade: false, solution: true -->
```python
# Split data into train and test. Keep 10% of samples for testing
# Divide the data into these variables - X_train, X_test, y_train, y_test


from sklearn.model_selection import train_test_split

# YOUR CODE HERE
```

<!-- grade_id: cell-4f278c6930c3ec50, points: 2, grade: true, solution: false -->
```python
"""test for splitting"""
```

### Question 5
#### Find mean and standard deviation of input data for normalization


<!-- grade_id: cell-07655e5eea8b4370, points: 0, grade: false, solution: true -->
```python
## Normalize the Data
def findMeanStddev(X):
    '''
    Input: 
        X: a matrix of size (no. of samples, dimension of each sample)
    Output:
        mean: mean of samples in X (same size as X)
        stddev: element-wise std dev of sample in X (same size as X)
    '''
    # YOUR CODE HERE
    
    return mean, stddev

```

<!-- grade_id: cell-644dd1a507ca5f16, points: 2, grade: true, solution: false -->
```python
"""test for findMeanStdev"""

```

### Question 6
#### Normalization

<!-- grade_id: cell-de210b0c0b54e506, points: 0, grade: false, solution: true -->
```python
def normalizeX(X, mean, stddev):
    '''
    Input:
        X: a matrix of size (no. of samples, dimension of each sample)
        mean: mean of samples in X (same size as X)
        stddev: element-wise std dev of sample in X (same size as X) 
    Output:
        Xn: X modified to have 0 mean and 1 std dev
    '''
    # YOUR CODE HERE
    
    return Xn

mean_train, stddev_train = findMeanStddev(X_train)
X_train = normalizeX(X_train, mean_train, stddev_train)
X_test = normalizeX(X_test, mean_train, stddev_train)
```

<!-- grade_id: cell-dff88f05fb986621, points: 3, grade: true, solution: false -->
```python
"""Test for Normalization"""

```

### Question 7
#### Creating and training the model

<!-- grade_id: cell-a5a12aabd0c3d59e, points: 0, grade: false, solution: true -->
```python
import keras
def create_model():
    """
    Create the NN model using Keras with single hidden layer of 10 neurons.
    Use sigmoid activation at the hidden layer and softmax activation at the output layer.
    Use the sgd optimizer and categorical_crossentropy for the loss.
    
    Inputs:
        None
    Outputs:
        model: keras model after compiling
    """
    # YOUR CODE HERE
    
    return model
  
model = create_model()
```

```python
# PROVIDE THE ANSWER HERE
# Eg if Total Parameters in the above model is 100 then write Total_Params = 100 below

Total_Params = 
```

<!-- grade_id: cell-98e8b63c2572431d, points: 5, grade: true, solution: false -->
```python
"""Test for create model"""
```

```python
## Create the Model and training (Code given)
model = create_model()
history = model.fit(X_train, y_train, epochs=15, batch_size = 10, validation_split = 0.1)
```

```python
## Plotting the accuracy with iterations (Code given)
from matplotlib import pyplot as plt
plt.plot(history.history['val_acc'])
```

<!-- grade_id: cell-d0bf3781ee836622, points: 6, grade: true, solution: false -->
```python
"""Test for model"""

```

## Part 3 - NLP


### Character-level Neural Language Model in Keras
A language model that predicts the next word in the sequence based on the specific words that have come before it in the sequence.

### Importing Libraries and Loading data
The nursery rhyme "Sing a Song of Sixpence" is well known in the west.

The first verse is common, but there is also a 4 verse version that we will use to develop our character-based language model.

```python
from numpy import array
from pickle import dump
from keras.utils import to_categorical
from keras.models import Sequential
from keras.layers import Dense
from keras.layers import LSTM

raw_text = """Sing a song of sixpence,
A pocket full of rye.
Four and twenty blackbirds,
Baked in a pie.
 
When the pie was opened
The birds began to sing;
Wasn't that a dainty dish,
To set before the king.
 
The king was in his counting house,
Counting out his money;
The queen was in the parlour,
Eating bread and honey.
 
The maid was in the garden,
Hanging out the clothes,
When down came a blackbird
And pecked off her nose."""
```

### Question 1
#### Preprocess

<!-- grade_id: cell-677facdf3167dca0, points: 0, grade: false, solution: true -->
```python
length = 10
def preprocess(raw_text,length):
    """
    Input:
        raw_text: a sequence of characters
          length: length of characters in 'sequences'(output) list 
    
    Action:
        
        1. Tokenize the raw_text
        2. Split the string in chunk of 'length' characters. For example, your output should look like sequences = ['Sing a song', 'ing a song ', 'ng a song o', ...]
        
    Output:
        tokens: a list obtained after splitting the raw_text
        sequences: a list obtained after splitting the 'raw_text' string in chunk of ('length') characters
    """
    # YOUR CODE HERE
    return sequences, tokens, raw_text
    
    
sequences, tokens, raw_text = preprocess(raw_text,length)
```

<!-- grade_id: cell-fa0e273fc35e3ae0, points: 3, grade: true, solution: false -->
```python
'''test for preprocess'''
```

### Question 2
####Mapping

<!-- grade_id: cell-31388153b6f2c37f, points: 0, grade: false, solution: true -->
```python
def mapp(raw_text):
    """
    Input:
        raw_text: 1d list
    
    Actions: 
        
        IMP: FOLLOW THE ACTIONS GIVEN BELOW IN THE PRESCRIBED SEQUENCES
        
        1. Create a set of raw_text
        2. Convert it into a list
        3. Sort the list and store it into a list (say chars) 
        
        Note: You can use inbuilt functions like set(), sorted(), list(), etc. 
        
        4. Create dictionary (say mapping) where each key is character present in sorted list 'chars' (Hint: use enumerate(chars))
        
    Output:
        mapping: dictionary
    """ 
    # YOUR CODE HERE
    
mapping = mapp(raw_text)
vocab_size = len(mapping)
```

<!-- grade_id: cell-9d51f92b9b734c27, points: 2, grade: true, solution: false -->
```python
'''test for mapp'''

```

### Question 3
####Encode Sequences

<!-- grade_id: cell-dd1ebeaad3a4fc3e, points: 0, grade: false, solution: true -->
```python
def encode(sequences,mapping):
    """
    Input:
        sequences: 1d list
        mapping : dictionary
    
    Action: 
        Encode all the sequences to corroponding dictionary mapping (integers) and store it in the 'sequences_encode' np array
        For E.g: sequences[0] = 'Sing a song'
                 sequences_encode[0] = [11 22 26 20  0 14  0 31 27 26 20]
                 
                 where S=11, i=22, n=26, etc. have been assigned to dictionary 'mapping'
        
        Important: data-type of sequences_encode should be numpy.ndarray.
    
    Output:
        sequences_encode: np 2d array
    """
    # YOUR CODE HERE
    
sequences_encode = encode(sequences,mapping)
```

<!-- grade_id: cell-84c0a9d1c9480f15, points: 3, grade: true, solution: false -->
```python
'''test for encode'''
def test_encode():
  assert sequences_encode.shape == (399,11)
  print('Test passed', '\U0001F44D')
test_encode()
```

###Question 4
####Split Input and Output

Now that the sequences have been integer encoded, we can separate the columns into input and output sequences of characters. <br>
First 10 characters will be our input and 11th character will be our output corrosponding to it.

<!-- grade_id: cell-d889309e55ddcf83, points: 0, grade: false, solution: true -->
```python
def split(sequences_encode):
    """
    Input:
        sequences_encode: np 2d array
        
    Action: Separate the columns into input and output sequences of characters. 
            First 10 characters will be our input (X) and 11th character will be our output (y) corrosponding to it.
            
    Output:
        X: np 2d array
        Y: np array
    """
    # YOUR CODE HERE
    
X, y = split(sequences_encode)
```

<!-- grade_id: cell-5f23b1b2b4fe44d5, points: 2, grade: true, solution: false -->
```python
'''test for split'''
def test_split():
  assert X.shape == (399,10)
  print('Test passed', '\U0001F44D')
test_split()
```

# SECTION II - MCQ/Short Answer test

## Part 1- RTRL (Recurrent Neural Network)


Suppose 'y' is a timeseries data (which is obtained from the following linear equation) 

####$y(t+1) = w_0y(t) + w_1y(t-1) + w_2u(t) $<br>

####$w = [w_0, \; w_1, \; w_2]$
####If  $\; w = [1.1, \; -0.3, \;0.5], \; y(t)=2, \; y(t-1) = 0.8 \;$ and $\; u(t) = 1.2$ <br><br>
$\Large \frac{\partial y(t)}{\partial w_0} = 0, \; \; \frac{\partial y(t-1)}{\partial w_0} = 1.5$<br><br>
$\Large \frac{\partial y(t)}{\partial w_1} = 0.7, \; \; \frac{\partial y(t-1)}{\partial w_1} = 0.2$<br><br>
$\Large \frac{\partial y(t)}{\partial w_2} = 1, \; \; \frac{\partial y(t-1)}{\partial w_2} = 0.6$



###Question 1. Which of the following is correct
<br>
$$
a) \frac{\partial y(t+1)}{\partial w_0} = w_0\frac{\partial y(t)}{\partial w_0} + y(t-1) + w_1\frac{\partial y(t-1)}{\partial w_0}
$$
<br>
$$
b) \frac{\partial y(t+1)}{\partial w_0} = w_0\frac{\partial y(t)}{\partial w_0} + y(t) + w_1\frac{\partial y(t-1)}{\partial w_0}
$$
<br>
$$
c) \frac{\partial y(t+1)}{\partial w_0} = w_0\frac{\partial y(t)}{\partial w_0} + w_oy(t) + w_1\frac{\partial y(t-1)}{\partial w_0}
$$
<br>
$$
d) \frac{\partial y(t+1)}{\partial w_0} = w_0\frac{\partial y(t)}{\partial w_0} + u(t) + w_1\frac{\partial y(t-1)}{\partial w_0}
$$
<br>
###Note: 
Write<br>
ans1a='a' in the code cell provided below (if answer is 'a')<br>
else ans1a='b' (if answer is 'b')<br>
else ans1a='c' (if answer is 'c')<br>
else ans1a='d' (if answer is 'd')<br>

```python
ans1a = ''
```

<!-- grade_id: cell-7ceb5c3b075dabf5, points: 2, grade: true, solution: false -->
```python
'''test for 1a'''
def test_1a():
  
test_1a()
```

###Question 2. Which of the following is correct

$$a) \; \large \frac{\partial y(t+1)}{\partial w_2} = 1.52$$<br>
$$b) \; \large \frac{\partial y(t+1)}{\partial w_2} = 2.62$$<br>
$$c) \; \large \frac{\partial y(t+1)}{\partial w_2} = 4.32$$<br>
$$d) \; \large \frac{\partial y(t+1)}{\partial w_2} = 2.12$$<br>

###Note: 
Write<br>
ans1b='a' in the code cell provided below (if answer is 'a')<br>
else ans1b='b' (if answer is 'b')<br>
else ans1b='c' (if answer is 'c')<br>
else ans1b='d' (if answer is 'd')<br>

```python
ans1b = ''
```

<!-- grade_id: cell-57146823cd6cc87c, points: 2, grade: true, solution: false -->
```python
'''test for 1b'''
def test_1b():
  
test_1b()
```

###Question 3. Which of the following is correct

$$a) \; \large \frac{\partial y(t+1)}{\partial w_1} = 1.52$$<br>
$$b) \; \large \frac{\partial y(t+1)}{\partial w_1} = 1.51$$<br>
$$c) \; \large \frac{\partial y(t+1)}{\partial w_1} = 2.32$$<br>
$$d) \; \large \frac{\partial y(t+1)}{\partial w_1} = 2.12$$<br>

###Note: 
Write<br>
ans1c='a' in the code cell provided below (if answer is 'a')<br>
else ans1c='b' (if answer is 'b')<br>
else ans1c='c' (if answer is 'c')<br>
else ans1c='d' (if answer is 'd')<br>

```python
ans1c = ''
```

<!-- grade_id: cell-6658dd0450ca114d, points: 2, grade: true, solution: false -->
```python
'''test for 1c'''
def test_1c():
test_1c()
```

###Question 4. Which of the following is correct

$$a) \; \large \frac{\partial y(t+1)}{\partial w_0} = 1.55$$<br>
$$b) \; \large \frac{\partial y(t+1)}{\partial w_0} = 1.5$$<br>
$$c) \; \large \frac{\partial y(t+1)}{\partial w_0} = 2.35$$<br>
$$d) \; \large \frac{\partial y(t+1)}{\partial w_0} = 1.125$$<br>

###Note: 
Write<br>
ans1d='a' in the code cell provided below (if answer is 'a')<br>
else ans1d='b' (if answer is 'b')<br>
else ans1d='c' (if answer is 'c')<br>
else ans1d='d' (if answer is 'd')<br>

```python
ans1d = ''
```

<!-- grade_id: cell-7fefec420e3adc28, points: 2, grade: true, solution: false -->
```python
'''test for 1d'''
def test_1d():
  
test_1d()
```

###Question 5. Which of the following is correct weight update formula

###given $\; E = \frac{1}{2}(y_d - y)^2$ <br>




$$a) \; \large w_{new} = w_{old} - \eta(y_d - y)\cdot \frac{\partial y(t)}{\partial w} $$<br>
$$b) \; \large w_{new} = w_{old} + \eta(y_d - y)\cdot \frac{\partial y(t+1)}{\partial w} $$<br>
$$c) \; \large w_{new} = w_{old} - \eta(y_d - y)\cdot \frac{\partial y(t+1)}{\partial w} $$<br>
$$d) \; \large w_{new} = w_{old} + \eta(y_d - y)\cdot \frac{\partial y(t)}{\partial w} $$<br>

###Note: 
Write<br>
ans1e='a' in the code cell provided below (if answer is 'a')<br>
else ans1e='b' (if answer is 'b')<br>
else ans1e='c' (if answer is 'c')<br>
else ans1e='d' (if answer is 'd')<br>

```python
ans1e = ''
```

<!-- grade_id: cell-d782350d0c1914f1, points: 2, grade: true, solution: false -->
```python
'''test for 1e'''
def test_1e():
test_1e()
```

###Question 6. If  $\; \eta = 1, \; y_d = 2.66\; and\; w_{old}=[1.1,\; -0.3,\; 0.5].\; $ Then, which of the following is correct <br><br>

$$a) \; \large w_{new} = [0.915, \; -0.461,\; 0.188]$$<br>
$$b) \; \large w_{new} = [0.965, \; -0.481,\; 0.388]$$<br>
$$c) \; \large w_{new} = [0.925, \; -0.411,\; 0.488]$$<br>
$$d) \; \large w_{new} = [0.945, \; -0.451,\; 0.288]$$<br>

###Note: 
Write<br>
ans1f='a' in the code cell provided below (if answer is 'a')<br>
else ans1f='b' (if answer is 'b')<br>
else ans1f='c' (if answer is 'c')<br>
else ans1f='d' (if answer is 'd')<br>

```python
ans1f = ''
```

<!-- grade_id: cell-a12f65a0f422a36e, points: 2, grade: true, solution: false -->
```python
'''test for 1f'''
def test_1f():
test_1f()
```

## Part 2 - CNN (Convolutional Neural Network)

### Question 1
In the following figure,what is the total number of tunable parameters including the biases?

<img src="https://drive.google.com/uc?id=1yqMW2TSlME3YrjlspOUbVxnayHws3-3v" title="" align="center" width="50%" height="100%"/> <br>


```python
# type the value of the total number of parameters
Num_Param = 
```

<!-- grade_id: cell-c340a3455b7beae5, points: 1, grade: true, solution: false -->
```python
'''test for Num_of_Param '''
```

### Question 2

The CNN used to solve MNIST hand written digit identification problem has following
architecture:
*   Input has the dimension 27x27x1 .
*   The first layer uses 25 filters of kernel size 5x5. The convolution takes place with stride = 1, Padding is zero and activation is ReLU
*   The second layer employs 30 filters of kernel size of 7x7, with stride = 2, padding is zero and activation is ReLU
*   The output layer is a simple classifier (fully connected) that actuates the actual digit in terms of one-hot coding.
####Find the dimension of the response in the first layer


```python
# Fill the dimensions below
dim_first = [ , , ]
```

<!-- grade_id: cell-ab56c262a5e6e553, points: 1, grade: true, solution: false -->
```python
'''test for dim_first'''
```

###Question 3
In Question 2 above, Find the dimension of the response in the second layer.

```python
# Fill the dimensions below
dim_second = [ , , ]
```

<!-- grade_id: cell-2e764271dcc35e3f, points: 1, grade: true, solution: false -->
```python
'''test for dim_second'''
```

###Question 4 
Find the number of parameters that will be estimated using gradient descent algorithm (including the biases)

```python
#Provide the value of num_grad
num_grad = 
```

<!-- grade_id: cell-225708842107d37e, points: 1, grade: true, solution: false -->
```python
'''test'''
```

## Part 3 - Keras Model Building

### Question 1
In the following python code using a sequential model from keras library  <br>


model.add(Dense(10, input_dim= 100, activation='relu')) <br>
model.add(Dense(8, activation='relu')) <br>
model.add(Dense(5, activation='relu')) <br>
model.add(Dense(5, activation='softmax')) <br>

#### What is the total number of weights to be trained (including the biases)? <br> 


```python
Num_of_Param = 
```

<!-- grade_id: cell-11069dbc71b45a88, points: 2, grade: true, solution: false -->
```python
'''test for Num_of_Param '''
```

```python

```